# BT01: Ẩn tin mật trên văn bản bằng phương pháp chèn khoảng trắng <font color="grey">(chèn khoảng trắng giữa các từ để làm cho văn bản được căn lề 2 bên)


Họ tên: Trịnh Văn Minh

MSSV: 1712601

---
## Cách làm bài và nộp bài
&#9889; Bạn lưu ý là mình sẽ dùng chương trình hỗ trợ chấm bài nên bạn cần phải tuân thủ chính xác qui định mà mình đặt ra, nếu không rõ thì hỏi, chứ không nên tự tiện làm theo ý của cá nhân.

**Cách làm bài**

Bạn sẽ làm trực tiếp trên file notebook này. Đầu tiên, bạn điền họ tên và MSSV vào phần đầu file ở bên trên. Trong file, bạn làm bài ở những chỗ có ghi là:
```python
# YOUR CODE HERE
raise NotImplementedError()
```
hoặc đối với những phần code không bắt buộc thì là:
```python
# YOUR CODE HERE (OPTION)
```
hoặc đối với markdown cell thì là:
```markdown
YOUR ANSWER HERE
```
Tất nhiên, khi làm thì bạn xóa dòng `raise NotImplementedError()` đi.
Đối những phần yêu cầu code thì thường ở ngay phía dưới sẽ có một (hoặc một số) cell chứa các bộ test để giúp bạn biết đã code đúng hay chưa; nếu chạy cell này không có lỗi gì thì có nghĩa là qua được các bộ test. Trong một số trường hợp, các bộ test có thể sẽ không đầy đủ; nghĩa là, nếu không qua được test thì là code sai, nhưng nếu qua được test thì chưa chắc đã đúng hoàn toàn.

Trong khi làm bài, bạn có thể cho in ra màn hình, tạo thêm các cell để test. Nhưng khi nộp bài thì bạn xóa các cell mà bạn tự tạo, xóa hoặc comment các câu lệnh in ra màn hình. Bạn lưu ý <font color=red>không được tự tiện xóa các cell hay sửa code của Thầy</font> (trừ những chỗ được phép sửa như đã nói ở trên).

Trong khi làm bài, thường xuyên `Ctrl + S` để lưu lại bài làm của bạn, tránh mất mát thông tin.


*Nên nhớ mục tiêu chính ở đây là <font color=green>học, học một cách chân thật</font>. Bạn có thể thảo luận ý tưởng với bạn khác, nhưng <font color=green>code và bài làm phải là của bạn, dựa trên sự hiểu thật sự của bạn</font>. <font color=red>Nếu vi phạm thì sẽ bị 0 điểm cho toàn bộ môn học.</font>*

**Cách nộp bài**

Khi chấm bài, đầu tiên mình sẽ chọn `Kernel` - `Restart Kernel & Run All Cells`, để restart và chạy tất cả các cell trong notebook của bạn; do đó, trước khi nộp bài, bạn nên chạy thử `Kernel` - `Restart Kernel & Run All Cells` để đảm bảo mọi chuyện diễn ra đúng như mong đợi.

Sau đó, bạn tạo thư mục nộp bài theo cấu trúc sau:
- Thư mục `MSSV` (vd, nếu bạn có MSSV là 1234567 thì bạn đặt tên thư mục là `1234567`)
    - Thư mục `BT01`
        - File `BT01-AnTinMatTrenVanBan_KhoangTrang.ipynb` (không cần nộp các file khác)

Cuối cùng, bạn nén thư mục `MSSV` này lại và nộp ở link trên moodle. <font color=red>Bạn lưu ý tuân thủ chính xác cấu trúc này.</font>

## 2. Import

```python
import bitarray
```

## 3. Hàm nhúng (6đ)

<!-- grade_id: cell-2f21b430087c5c35, points: 0, grade: false, solution: true -->
```python
def embed(msg_file, cover_text_file, text_width, stego_text_file):
    '''
    Nhúng tin mật vào văn bản bằng phương pháp chèn khoảng trắng 
    để văn bản được căn lề 2 bên.
    
    Các tham số:
        msg_file (str): Tên file chứa secret message.
        cover_text_file (str): Tên file chứa cover text.
        text_width (int): Chiều dài của dòng sau khi được căn lề.
        stego_text_file (str): Tên file chứa stego text.
    Giá trị trả về:
        bool: True nếu nhúng thành không, False nếu không đủ chỗ để nhúng. 
    '''
    # YOUR CODE HERE
    def get_msg_bits(msg_file):
        with open(msg_file,'rb') as mf:
            msg_bytes = mf.read()
        msg_bits = bitarray.bitarray()
        msg_bits.frombytes(msg_bytes)
        return msg_bits

    def get_cover_array(cover_text_file):
        cover_array = []
        with open(cover_text_file,'r') as ctf:
            for line in ctf:
                cover_array.append(line.replace('\n',''))
        return cover_array

    def counting_space(msg):
        n = msg.count(' ')
        return n

    def is_alignable(line, next_line, text_width):
        if next_line == '':
            return False
        next_word = next_line.split()[0]
        n = counting_space(line)
        m = text_width - len(line)
        if n < 1 or m > len(next_word):
            return False
        return True

    def alignable_array(cover_line_array, text_width):
        len_line = len(cover_line_array)
        j = []
        for i in range(len_line - 1):
            j.append(is_alignable(cover_line_array[i], cover_line_array[i + 1], text_width))
        j.append(False)
        return j

    def bit2space(bit):
        #bit 0
        if not bit:
            return ['  ',' ']
        # bit 1
        return [' ','  ']

    def bits2space_array(msg_bits):
        embeded_space_array = []
        for i in msg_bits:
            embeded_space_array.append(bit2space(i))

        return embeded_space_array

    def embedding(cover_line_array, ali_array, stego_spaces_linearray):

        embeded_line_array = []
        # stego_spaces_linearray is a lines array contain spaces that was embeded
        # j is the index of stego_spaces_linearray
        j = 0
        #len cover_line_array = len ali_array - mảng xác định khả năng căn lề true or false
        for i in range(len(ali_array)):
            if ali_array[i]:
                stego_line_msg = ''
                # stego_line format ['  ', ' ', ' ', '  ', '  ', ' ', '  ', ' ', ' ', ' ']
                stego_line = stego_spaces_linearray[j]
                # 'Have you heard about steganography? I find it very interesting and'
                cover_line = cover_line_array[i].split()


                for i in zip(cover_line, stego_line):
                    stego_line_msg += i[0]+i[1]
                stego_line_msg += cover_line[-1]

                embeded_line_array.append(stego_line_msg)
                j += 1
            else:
                # cover_line_array is a lines array of cover text
                embeded_line_array.append(cover_line_array[i])

        return embeded_line_array
        
    def save_stego_file(embeded_line_array, stego_text_file):
        with open(stego_text_file,'w') as stf:
            for i in range(len(embeded_line_array) - 1):
                stf.write(embeded_line_array[i] + '\n')
            stf.write(embeded_line_array[-1])

    b = 0 # Chỉ số duyệt từng bit của msg_bits
    l = 0 # Chỉ số duyệt từng dòng của cover_line_array
    
    '''
    flag =  -1: flase,
             0: đã chèn vào đuôi được bit số 1,
             1: đã chèn vào đuôi được 10.....
     '''
    
    flag = -1 
    msg_bits = get_msg_bits(msg_file)
    msg_spaces = bits2space_array(msg_bits)

    # -------
    stego_spaces_linearray = []
    # -------

    #một mảng gồm các dòng của cover_text
    cover_line_array = get_cover_array(cover_text_file)

    ali_array = alignable_array(cover_line_array,text_width)
    # Lặp trong khi mà l < số dòng của cover_line_array

    while(l < len(cover_line_array)):
        # xác định dòng thứ l có phải dòng căn lề
        if(ali_array[l]):           
        
            n = counting_space(cover_line_array[l])
            m = text_width - len(cover_line_array[l])

            if n > m and m > 0:
                # Nhúng được min(m, n-m) bit
                # Duyệt min(m, n-m) cặp khoảng trắng đầu của dòng l
                stego_spaces_line = []
                
                able_bits = min(m, n - m)
                for _ in range(able_bits):
                    if b < len(msg_bits):
                        stego_spaces_line.extend(msg_spaces[b])
                        b+=1
                    else:
                        if flag == -1:
                            #if the first 'else' --> embed bit 1
                            stego_spaces_line.extend([' ','  '])
                            #chèn thành công số 1, thay đổi flag để k chạy vào lần sau
                            flag = 0
                        else:
                            #chèn thành công số 0, thay đỔi flag thành True
                            flag = 1
                            stego_spaces_line.extend(['  ',' '])
                if able_bits == (n - m):
                    spaces1_remain = (2*m-n)*['  ']  
                    stego_spaces_line.extend(spaces1_remain)
                
                bk = len(stego_spaces_line)
                if bk < n:
                    stego_spaces_line.extend(' '*(n - bk))


                stego_spaces_linearray.append(stego_spaces_line)                
            else:
                # m >= n
                # không nhúng được bit, nhưng mỗi khoảng trắng cần được chèn thêm ít nhất một khoảng trắng để dòng l được căn lề
                spaces_divided = (m // n) + 1
                spaces_remain = m % n
                stego_spaces_line = []
               
                stego_spaces_line.extend([(spaces_divided + spaces_remain)*' '])
                stego_spaces_line.extend((n-1)*[spaces_divided*' '])
                stego_spaces_linearray.append(stego_spaces_line)   

            l = l + 1 

        else:
            l = l + 1
    
    # endwhile --> Nếu vẫn chưa nhúng được bit 1 (trong đuôi 100...): NHÚNG THẤT BẠI!

    # flag == 1 --> nhúng thành công đuôi 10......
    if flag == 1:
        embeded_line_array = embedding(cover_line_array, ali_array, stego_spaces_linearray)
        save_stego_file(embeded_line_array,stego_text_file)
        return True

    return False
```

<!-- grade_id: cell-057a58bf038e1e60, points: 5, grade: true, solution: false -->
```python
# TEST 1            
result = embed('msg.txt', 'cover.txt', 70, 'stego.txt')
assert result == True
with open('stego.txt', 'r') as f:
    stego_text = f.read()
with open('correct_stego.txt', 'r') as f:
    correct_stego_text = f.read()
assert stego_text == correct_stego_text
```

<!-- grade_id: cell-f625774497df6542, points: 1, grade: true, solution: false -->
```python
# TEST 2
result = embed('msg2.txt', 'cover.txt', 70, 'stego.txt')
assert result == False
```

## 4. Hàm rút trích (4đ)

<!-- grade_id: cell-f3f9b34770f282aa, points: 0, grade: false, solution: true -->
```python
def extract(stego_text_file, extr_msg_file):
    '''
    Hàm rút trích tin mật đã được nhúng bằng phương pháp chèn khoảng trắng.
    
    Các tham số:
        stego_text_file (str): Tên file chứa stego text.
        extr_msg_file (str): Tên file chứa secret message được rút trích.
    '''
    # YOUR CODE HERE
    import re
    
    stego_text = []
    with open(stego_text_file, 'r') as f:
        for line in f:
            stego_text.append(line.replace('\n',''))


    extracted_msg_bits = []
    l = 0 # Chỉ số duyệt từng dòng của stego_text
    while(l < len(stego_text)):
        #nếu là dòng căn lề
        if len(stego_text[l]) == 70:
            #  Duyệt các khoảng trắng của dòng l
            # ['  ', ' ', ' ', '  ', '  ', ' ', '  ', ' ', ' ', ' ']
            spaces = re.findall(' +', stego_text[l])
            for i in range(1, len(spaces), 2):
                if spaces[i - 1] == '  ' and spaces[i] == ' ':
                    extracted_msg_bits.extend([False])
                elif spaces[i - 1] == ' ' and spaces[i] == '  ':
                    extracted_msg_bits.extend([True])
                else:
                    break
            l+=1
        else:
            l+=1

    # Cắt đoạn đuôi 100.. ra khỏi extracted_msg_bits
    
    i = -1
    while(extracted_msg_bits[i] == False):
        i -= 1
    list_msg_bits = extracted_msg_bits[:i]

    #save to extr_msg_file

    # [False, True, False,.....] --> bitarray('010.....')
    msg_bits = bitarray.bitarray(list_msg_bits)
    # bit to bytes to msg
    msg = msg_bits.tobytes().decode('utf-8')
    #save
    with open(extr_msg_file,'w') as f:
        f.write(msg)
    
```

<!-- grade_id: cell-caf1afe26d95b670, points: 4, grade: true, solution: false -->
```python
# TEST
extract('correct_stego.txt', 'extr_msg.txt')
with open('extr_msg.txt', 'r') as f:
    extr_msg = f.read()
with open('correct_extr_msg.txt', 'r') as f:
    correct_extr_msg = f.read()
assert extr_msg == correct_extr_msg
```